```scala

```

# CSCI 3155: Assignment 6 (45 points)

Topics Covered: Function calls in Lettuce

__Name__: NICOLE DONG

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
passed: (points: Int)Unit
```

## Problem 1 (45 Points): Multiple Argument Functions

In class, we have explored just single argument functions in lettuce. In this problem, we will 
explore, multiple argument functions.

Consider: 

~~~
let foo1 = function (x, y)
            x - 2 * y
          in 
       foo1(10, 15)
~~~

this code should return -20

We will allow zero arguments as well.

~~~
let x = 5 in 
let bar1 = function()
           x 
           in 
    bar1()
~~~

This code should return 5.

## 1 A (7 points): Extending the Abstract Syntax

Consider the grammar specification we have seen thus far.

$$\begin{array}{rcll}
\mathbf{Program} & \rightarrow & TopLevel(\mathbf{Expr}) \\[5pt]
\mathbf{Expr} & \rightarrow & Const(\mathbf{Number}) \\
 & | & Ident(\mathbf{Identifier}) \\
 & | & Plus(\mathbf{Expr}, \mathbf{Expr}) \\
 & | & FunDef( \mathbf{Identifier}^*, \mathbf{Expr}) & \text{Note multiple parameters possible} \\ 
 & | & FunCall(\mathbf{Expr}, \mathbf{Expr}^*) & \text{function call - expr(expr1, ... , exprn)} \\
 & | & Let(\mathbf{Identifier},\mathbf{Expr}, \mathbf{Expr})  \\
\end{array}$$

Write the scala definition for `FunDef` and `FunCall`. Please use lists to implement the Kleene Star.

<!-- grade_id: cell-925e2ee1e5248e66, points: 0, grade: false, solution: true -->
```scala
sealed trait Program
sealed trait Expr
case class Const(f: Double) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Let(x: String, e1: Expr, e2: Expr) extends Expr
case class FunDef(l: List[String], e: Expr) extends Expr
case class FunCall(e: Expr, l: List[Expr]) extends Expr
case class TopLevel(e: Expr) extends Program
```

Output:
```
defined trait Program
defined trait Expr
defined class Const
defined class Ident
defined class Plus
defined class Let
defined class FunDef
defined class FunCall
defined class TopLevel
```

<!-- grade_id: cell-0a323348d10eb569, points: 3, grade: true, solution: false -->
```scala
//BEGIN TEST
// TEST Two arguments
val x = Ident("x")
val y = Ident("y")
val foo = Ident("foo")
val fun1 = FunDef(List("x", "y"), Plus(x, y))
val l1 = Let("foo", fun1, FunCall(foo, List(Const(10.0), Const(20.0))))
val p1 = TopLevel(l1)
passed(3)
//END TEST
```

Output:
```
*** Tests Passed (3 points) ***
x: Ident = Ident(x)
y: Ident = Ident(y)
foo: Ident = Ident(foo)
fun1: FunDef = FunDef(List(x, y),Plus(Ident(x),Ident(y)))
l1: Let = Let(foo,FunDef(List(x, y),Plus(Ident(x),Ident(y))),FunCall(Ident(foo),List(Const(10.0), Const(20.0))))
p1: TopLevel = TopLevel(Let(foo,FunDef(List(x, y),Plus(Ident(x),Ident(y))),FunCall(Ident(foo),List(Const(10.0), Const(20.0)))))
```

<!-- grade_id: cell-ac281039a6b38cb1, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
//Test zero arguments
val x = Ident("x")
val foo = Ident("foo")
val fcall = FunCall(foo, List())
val fdef = FunDef(List(), x)
val l2 = Let("foo", fdef, fcall)
val l3 = Let("x", Const(10), l2)
val p2 = TopLevel(l3)
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
x: Ident = Ident(x)
foo: Ident = Ident(foo)
fcall: FunCall = FunCall(Ident(foo),List())
fdef: FunDef = FunDef(List(),Ident(x))
l2: Let = Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List()))
l3: Let = Let(x,Const(10.0),Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List())))
p2: TopLevel = TopLevel(Let(x,Const(10.0),Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List()))))
```

## 1B (8 points): Define Closure for Function Calls with Zero or More Args.

We will now redefine closures for functions with zero or more args. When our function had one argument, our closures were defined as `Closure(id, expr, env)`. We would now like to define closures as
`Closure( [id1, ..., idn], expr, env)` where 
  - `id1..., idn` are the list of arguments for the function to be called. 
  - `expr` is the body of the function and 
  - `env` is the stored environment for static scoping.
 
 $$\begin{array}{rcl}
 \mathbf{Value} & \Rightarrow & num(\mathbf{Double}) \\
 & \Rightarrow & closure(\mathbf{String}^*, \mathbf{Expr}, \mathbf{Environment}) \\
 & \Rightarrow & error \\
 \end{array}$$
 
 For __Environment__ please use a scala immutable map from __String__ to __Value__

<!-- grade_id: cell-60259c7b459335ef, points: 0, grade: false, solution: true -->
```scala
sealed trait Value
case class Num(d: Double) extends Value
case object Error extends Value
case class Closure(l: List[String], e: Expr, env: Map[String, Value]) extends Value
```

Output:
```
defined trait Value
defined class Num
defined object Error
defined class Closure
```

<!-- grade_id: cell-599609ca6ef74bb3, points: 8, grade: true, solution: false -->
```scala
//BEGIN TEST
val v1 = Closure(List("x", "y", "z"), Plus(x, y), Map.empty)
val v2 = Closure(Nil, Const(15), Map.empty)
val env1: Map[String, Value] = Map("x" -> Num(25.0), "y"-> v2)
val v3 = Closure(List("x"), Plus(x, Const(25)), env1 )
passed(8)
//END TEST
```

Output:
```
*** Tests Passed (8 points) ***
v1: Closure = Closure(List(x, y, z),Plus(Ident(x),Ident(y)),Map())
v2: Closure = Closure(List(),Const(15.0),Map())
env1: Map[String,Value] = Map(x -> Num(25.0), y -> Closure(List(),Const(15.0),Map()))
v3: Closure = Closure(List(x),Plus(Ident(x),Const(25.0)),Map(x -> Num(25.0), y -> Closure(List(),Const(15.0),Map())))
```

## 1 C (30 Points) Build an Interpreter

Build an interpreter using the following semantic rules. Ensure that your interpreter correctly deals with the cases that give rise to error by throwing an exception.

$$\newcommand\semRule[3]{\begin{array}{c} #1 \\ \hline #2 \\\end{array} (\text{#3})} $$
$$\newcommand\eval{\mathbf{eval}}$$
$$\semRule{}{\eval(\texttt{FunDef([id1,..., idk], e)},\sigma) = \text{Closure}(\texttt{[id1,..., idk]}, \texttt{e}, \sigma)}{fundef}$$

$$\semRule {\eval(\texttt{e}, \sigma) = \text{Closure}(\texttt{[id1,..., idn]}, \texttt{fBody}, \color{red}{\sigma_{cl}}),\  \color{red}{n = k},\ (\forall\ i \in \{ 1, \ldots, k\})\ \eval(\texttt{ei}, \sigma) = v_i,  v_i \not= \mathbf{error}}
           {\eval(\texttt{FunCall(e, [e1, ..., ek])}, \sigma) = \eval(\texttt{fBody}, \color{red}{\sigma_{cl} \circ [id1 \rightarrow v_1, \ldots, idk \rightarrow v_k]})}{funcall-ok}$$
           
Let us interpret this rule. Read the points below carefully.
   - __Purpose__ : Evaluate an expression of the form `FunCall(e, [e1,...,ek])` where `e` is the expr for the called function, and `e1, ..., ek` are exprs for the arguments of this call. There are $k$ arguments.
     - (A) Evaluating `e` must yield a closure of the form $\text{Closure}(\texttt{[id1,..., idn]}, \texttt{fBody}, \color{red}{\sigma_{cl}})$.
     - (B) The number of arguments for the closure $n$ must equal that of the function call $k$.
     - (C) Evaluating each of the $k$ arguments `ei` for $i = 1, \ldots, k$ must yield $v_i$ where $v_i$ is not error.
     - (D) Then the result of evaluating the call is the same as that of evaluating `fBody` under the environment $\color{red}{\sigma_{cl}}$ extended by mapping the formal parameters `id1.., idk` to $v_1, \ldots, v_k$, respectively.


We can write some error rules.

$$\semRule {\eval(\texttt{e}, \sigma) \not\in \text{Closure}}
           {\eval(\texttt{FunCall(e, [e1, ..., ek])}, \sigma) = \mathbf{error}}{funcall-not-a-function}$$


$$\semRule {\eval(\texttt{e}, \sigma) = \text{Closure}(\texttt{[id1,..., idn]}, \texttt{fBody}, \color{red}{\sigma_{cl}}),\  \color{red}{n \not= k}}
           {\eval(\texttt{FunCall(e, [e1, ..., ek])}, \sigma) = \mathbf{error}}{funcall-wrong-num-args}$$
           
  
$$\semRule {\eval(\texttt{e}, \sigma) = \text{Closure}(\texttt{[id1,..., idn]}, \texttt{fBody}, \color{red}{\sigma_{cl}}),\  \color{red}{n = k},\ (\exists\ i \in \{1, \ldots, k\})\ \eval(\texttt{ei}, \sigma) = \mathbf{error}}
           {\eval(\texttt{FunCall(e, [e1, ..., ek])}, \sigma) = \mathbf{error}}{funcall-arg-error}$$
           
           
**Important:** The use of `var` and `while/for` loop is not allowed. You can use the functions map, foldLeft and zip over lists.

**Hint** The following function will be very useful for you. Try it out and see what it does!

~~~
val l1 = List(10, 15, 20, 25, 30)
val l2 = List("a", "b", "c", "d", "e")
val l3 = l1.zip(l2)
println(l3)
~~~
 

<!-- grade_id: cell-cc0ebc1ead5d46dd, points: 0, grade: false, solution: true -->
```scala
class ErrorException(s:String) extends Exception(s){}  


def valueToNumber(v: Value): Double = v match {
    case Num(d) => d
    case _ => throw new ErrorException(s"Could not convert value $v to a number")
}

// get function body from Closure
def getfBody(x: Value): Expr = x match{
    case Closure(l, e, env) => {
        return e
    }
    case _ => {
        throw new IllegalArgumentException("Error")
    }
}

// get idList from Closure
def getid(x: Value): List[String] = x match{
    case Closure(l, e, env) => {
        return l
    }
    case _ => {
        throw new IllegalArgumentException("Error")
    }
}

// get env from Closure
def getEnv(x: Value): Map[String, Value] = x match{
    case Closure(l, e, env) => {
        return env
    }
    case _ => {
        throw new IllegalArgumentException("Error")
    }
}

// convert Expr to string for e
def getExpr(x: Expr): String = x match{
    case Ident(x) => {
        return x
    }
    case _ => {
        throw new IllegalArgumentException("Error")
    }
}

// check if eval(e, env) returns a Closure
def checkExpr(x: Value): Boolean = x match{
    case Closure(l, e, env) => {
        return true
    }
    case _ => {
        return false
    }
}

def eval(e: Expr, env: Map[String, Value]): Value = {
    def addValues(v1: Value, v2: Value): Value = 
        Num ( valueToNumber(v1) + valueToNumber(v2) )
    
    e match {
        case Const(d) => Num(d)
        case Ident(x) => {
            // your code here
            if(env.contains(x)){
                return env(x)
            }
            else{
                return Error
            }
        }
        case Plus(e1,e2) => {
            val v1 = eval(e1, env)
            val v2 = eval(e2, env)
            return addValues(v1,v2)
        }
        case Let(id, e1, e2) => {
            // your code here
            val s1 = eval(e1, env)
            val envtemp = collection.mutable.Map() ++ env
            envtemp(id) = s1
            val envnew = collection.immutable.Map() ++ envtemp
            return eval(e2, envnew)
        }
        
        case FunDef(idList, e) => {
            return Closure(idList, e, env)
        }
        
        case FunCall(e, eList) => {
            
            //println("e = " + e)
            //println("eList = " + eList)
            //println("env = " + env)
            
            // error: not a function
            if(!checkExpr(eval(e, env))){
                return Error
            }

            // Get the environment from Closure
            val env_cl = getEnv(eval(e, env))
            
            // Get the function body from Closure
            val fBody = getfBody(eval(e, env))
            
            // create a map using the idList from Closure and the eList
            val idList = getid(eval(e, env))
            val vList = eList.map(ei => eval(ei, env))
            
            // error: arg error
            if(!vList.filter(_ == Error).isEmpty){
                return Error
            }
            
            // error: wrong num args
            if(idList.length != vList.length){
                return Error
            }
            
            // combine the new map and the map from Closure
            return eval(fBody, env_cl ++ ((idList zip vList) toMap))
        }
    }
}


def evalProgram(p: Program): Value = p match {
    case TopLevel(e) => try
            eval(e, Map.empty)
    catch {
        case e: ErrorException => {
            println(e)
            Error
        }
        case e: IllegalArgumentException => {
            println(e)
            Error
        }
        case e => {
            println("Unknown Exception " + e.toString)
            Error
        }
            
    }
}
```

Output:
```
<console>:181: warning: This catches all Throwables. If this is really intended, use `case e : Throwable` to clear this warning.
               case e => {
                    ^
warning: there was one feature warning; re-run with -feature for details
defined class ErrorException
valueToNumber: (v: Value)Double
getfBody: (x: Value)Expr
getid: (x: Value)List[String]
getEnv: (x: Value)Map[String,Value]
getExpr: (x: Expr)String
checkExpr: (x: Value)Boolean
eval: (e: Expr, env: Map[String,Value])Value
evalProgram: (p: Program)Value
```

<!-- grade_id: cell-23be325bb4b874b2, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
// TEST Two arguments
val x = Ident("x")
val y = Ident("y")
val foo = Ident("foo")
// function(x, y)  x + y
val fun1 = FunDef(List("x", "y"), Plus(x, y))
//let foo = function(x,y) x+ y in foo(10, 20)
val l1 = Let("foo", fun1, FunCall(foo, List(Const(10.0), Const(20.0))))
//Program
val p1 = TopLevel(l1)
val v1 = evalProgram(p1)
println(s"Your program evaluated to $v1")
assert(v1 == Num(30.0), "Test 1 Failed")
passed(5)
//END TEST
```

Output:
```
Your program evaluated to Num(30.0)

*** Tests Passed (5 points) ***
x: Ident = Ident(x)
y: Ident = Ident(y)
foo: Ident = Ident(foo)
fun1: FunDef = FunDef(List(x, y),Plus(Ident(x),Ident(y)))
l1: Let = Let(foo,FunDef(List(x, y),Plus(Ident(x),Ident(y))),FunCall(Ident(foo),List(Const(10.0), Const(20.0))))
p1: TopLevel = TopLevel(Let(foo,FunDef(List(x, y),Plus(Ident(x),Ident(y))),FunCall(Ident(foo),List(Const(10.0), Const(20.0)))))
v1: Value = Num(30.0)
```

<!-- grade_id: cell-d59f51be3ed29a31, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
// TEST three arguments
val x = Ident("x")
val y = Ident("y")
val z = Ident("z")
val foo = Ident("foo2")
// function(x, y,z)  x + y+z+z
val fun2 = FunDef(List("x", "y", "z"), Plus(Plus(Plus(x, y), z), z))
//let foo = function(x,y) x+ y in foo(10, 20, 30)
val l2 = Let("foo2", fun2, FunCall(foo, List(Const(10.0), Const(20.0), Const(30.0))))
//Program
val p2 = TopLevel(l2)
val v2 = evalProgram(p2)
println(s"Your program evaluated to $v2")
assert(v2 == Num(90.0), "Test 2 Failed")
passed(5)
//END TEST
```

Output:
```
Your program evaluated to Num(90.0)

*** Tests Passed (5 points) ***
x: Ident = Ident(x)
y: Ident = Ident(y)
z: Ident = Ident(z)
foo: Ident = Ident(foo2)
fun2: FunDef = FunDef(List(x, y, z),Plus(Plus(Plus(Ident(x),Ident(y)),Ident(z)),Ident(z)))
l2: Let = Let(foo2,FunDef(List(x, y, z),Plus(Plus(Plus(Ident(x),Ident(y)),Ident(z)),Ident(z))),FunCall(Ident(foo2),List(Const(10.0), Const(20.0), Const(30.0))))
p2: TopLevel = TopLevel(Let(foo2,FunDef(List(x, y, z),Plus(Plus(Plus(Ident(x),Ident(y)),Ident(z)),Ident(z))),FunCall(Ident(foo2),List(Const(10.0), Const(20.0), Const(30.0)))))
v2: Value = Num(90.0)
```

<!-- grade_id: cell-8b5aa13d74ba3278, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
//Test zero arguments
val x = Ident("x")
val foo = Ident("foo")
val fcall = FunCall(foo, List())
val fdef = FunDef(List(), x)
val l2 = Let("foo", fdef, fcall)
val l3 = Let("x", Const(10), l2)
val p2 = TopLevel(l3)
val v3 = evalProgram(p2)
println(s"Your program evaluated to $v3")
assert(v3 == Num(10.0), "Test 3 Failed")
passed(5)
//END TEST
```

Output:
```
Your program evaluated to Num(10.0)

*** Tests Passed (5 points) ***
x: Ident = Ident(x)
foo: Ident = Ident(foo)
fcall: FunCall = FunCall(Ident(foo),List())
fdef: FunDef = FunDef(List(),Ident(x))
l2: Let = Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List()))
l3: Let = Let(x,Const(10.0),Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List())))
p2: TopLevel = TopLevel(Let(x,Const(10.0),Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List()))))
v3: Value = Num(10.0)
```

<!-- grade_id: cell-e305341e54d2e37c, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
//Evaluate with wrong number of args
val x = Ident("x")
val foo = Ident("foo")
val fcall = FunCall(foo, List(Const(1.0)))
val fdef = FunDef(List(), x)
val l2 = Let("foo", fdef, fcall)
val l3 = Let("x", Const(10), l2)
val p4 = TopLevel(l3)
val v4 = evalProgram(p4)
assert(v4 == Error, "Test 4 failed -- your program should have detected that arguments were mismatched")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
x: Ident = Ident(x)
foo: Ident = Ident(foo)
fcall: FunCall = FunCall(Ident(foo),List(Const(1.0)))
fdef: FunDef = FunDef(List(),Ident(x))
l2: Let = Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List(Const(1.0))))
l3: Let = Let(x,Const(10.0),Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List(Const(1.0)))))
p4: TopLevel = TopLevel(Let(x,Const(10.0),Let(foo,FunDef(List(),Ident(x)),FunCall(Ident(foo),List(Const(1.0))))))
v4: Value = Error
```

<!-- grade_id: cell-90f1081c0f420045, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
// TEST scoping
val x = Ident("x")
val y = Ident("y")
val z = Ident("z")
val foo = Ident("foo")

// function(x, y)  x + y + x + z 
val fun = FunDef(List("x", "y"), Plus(Plus(Plus(x, y), x),z))
val funcall = FunCall(foo, List(Const(10.0), Const(30.0)))
// let x = 1 in
//     let z = 50 in
//       let foo = fun(x,y) x + y + x + z in
//         let = 2 in
//           f(10,30)
val l4 = Let("x", Const(1.0), Let("z", Const(50.0), Let("foo", fun, Let("y", Const(2.0), funcall))))
//Program
val p2 = TopLevel(l4)
val v2 = evalProgram(p2)
println(s"Your program evaluated to $v2")
assert(v2 == Num(100.0), "Test 5 Failed == you did not evaluate static scoping correctly")
passed(5)
//END TEST
```

Output:
```
Your program evaluated to Num(100.0)

*** Tests Passed (5 points) ***
x: Ident = Ident(x)
y: Ident = Ident(y)
z: Ident = Ident(z)
foo: Ident = Ident(foo)
fun: FunDef = FunDef(List(x, y),Plus(Plus(Plus(Ident(x),Ident(y)),Ident(x)),Ident(z)))
funcall: FunCall = FunCall(Ident(foo),List(Const(10.0), Const(30.0)))
l4: Let = Let(x,Const(1.0),Let(z,Const(50.0),Let(foo,FunDef(List(x, y),Plus(Plus(Plus(Ident(x),Ident(y)),Ident(x)),Ident(z))),Let(y,Const(2.0),FunCall(Ident(foo),List(Const(10.0), Const(30.0)))))))
p2: TopLevel = TopLevel(Let(x,Const(1.0),Let(z,Const(50.0),Let(foo,FunDef(List(x, y),Plus(Plus(Plus(Ident(x),Ident(y)),Ident(x)),Ident(z))),Let(y,Const(2.0),FunCall(Ident(foo),List(Const(10.0), Const(30.0))))))))
v2: Value = Num(100.0)
```

<!-- grade_id: cell-4fe0ae3d93422dbe, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
//Evaluate with non function call
val x = Ident("x")
val fcall = FunCall(x, List(Const(1.0), Const(5.0), Const(20.0)))
val l3 = Let("x", Const(10), fcall)
val p5 = TopLevel(l3)
val v5 = evalProgram(p5)
assert(v5 == Error, "Test 6 failed -- your eval should have detected that a function call is made on a expression that is not a function")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
x: Ident = Ident(x)
fcall: FunCall = FunCall(Ident(x),List(Const(1.0), Const(5.0), Const(20.0)))
l3: Let = Let(x,Const(10.0),FunCall(Ident(x),List(Const(1.0), Const(5.0), Const(20.0))))
p5: TopLevel = TopLevel(Let(x,Const(10.0),FunCall(Ident(x),List(Const(1.0), Const(5.0), Const(20.0)))))
v5: Value = Error
```

```scala

```